---
title: Introduction
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: ca6ed460a9
xp: 50
```

`@projector_key`
968d548a1a11be2206d0ac11fd0bcbec

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: eea314ebef
xp: 50
```

Select two and four.

`@hint`
Find the correct answer

`@possible_answers`
one
[two]
three
[four]
five
six

`@feedback`
